# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which of the following are not part of the plotting capabilities of Python.

> A) Pandas
>
> B) Matplotlib
>
> C) BeautifulSoup
>
> D) Seaborn


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    ### BEGIN SOLUTION
    answer = "C"
    ### END SOLUTION
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
### BEGIN HIDDEN TESTS
assert q01() == "C"
### END HIDDEN TESTS
```

## #2 - In order to have plots display in your Jupyter notebook, which command do you need to run?

> A) `%matplotlib inline`
>
> B) `import plotly`
>
> C) `import BeautifulSoup`
>
> D) `%showplots`


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    ### BEGIN SOLUTION
    answer = "A"
    ### END SOLUTION
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
### BEGIN HIDDEN TESTS
assert q02() == "A"
### END HIDDEN TESTS
```

## #3 - Below is a explanation of what a chart can be used for.  Choose the best type of plot (barplot, histogram, scatterplot) to use.  Put your answer into a dictionary like this:

```
{
    'A': 'barplot',
    'B': 'histogram',
    'C': 'scatterplot',
    'D': 'histogram'
}
```

> A) Showing the relationship between two numerical values.
>
> B) Visualizing the distribution of a particular numerical value in a data set.
>
> C) Comparing a single statistic across different categories.
>
> D) Showing how the range and distribution of a numerical value differs between multiple data series.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    ### BEGIN SOLUTION
    answer = {
    'A': 'scatterplot',
    'B': 'histogram',
    'C': 'barplot',
    'D': 'histogram'
    }
    ### END SOLUTION
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == dict
### BEGIN HIDDEN TESTS
assert q03() == {
    'A': 'scatterplot',
    'B': 'histogram',
    'C': 'barplot',
    'D': 'histogram'
}
### END HIDDEN TESTS
```

```python

```